Acquisition phase-specific contribution of climbing fiber transmission to cerebellum-dependent motor memory

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public Review):

Summary:

The study by Seo et al highlights knowledge gaps regarding the role of cerebellar complex spike (CS) activity during different phases of learning related to optokinetic reflex (OKR) in mice. The novelty of the approach is twofold: first, specifically perturbing the activity of climbing fibers (CFs) in the flocculus (as opposed to disrupting communication between the inferior olive (IO) and its cerebellar targets globally); and second, examining whether disruption of the CS activity during the putative "consolidation phase" following training affects OKR performance.

The first part of the results provides adequate evidence supporting the notion that optogenetic disruption of normal CF-Purkinje neuron (PN) signaling results in the degradation of OKR performance. As no effects are seen in OKR performance in animals subjected to optogenetic irradiation during the memory consolidation or retrieval phases, the authors conclude that CF function is not essential beyond memory acquisition. However, the manuscript does not provide a sufficiently solid demonstration that their long-term activity manipulation of CF activity is effective, thus undermining the confidence of the conclusions.

Strengths:

The main strength of the work is the aim to examine the specific involvement of the CF activity in the flocculus during distinct phases of learning. This is a challenging goal, due to the technical challenges related to the anatomical location of the flocculus as well as the IO. These obstacles are counterbalanced by the use of a well-established and easy-to-analyse behavioral model (OKR), that can lead to fundamental insights regarding the long-term cerebellar learning process.

Weaknesses:

The impact of the work is diminished by several methodological shortcomings.

Most importantly, the key finding that prolonged optogenetic inhibition of CFs (for 30 min to 6 hours after the training period) must be complemented by the demonstration that the manipulation maintains its efficacy. In its current form, the authors only show inhibition by short-term optogenetic irradiation in the context of electrical-stimulation-evoked CSs in an ex vivo preparation. As the inhibitory effect of even the eNpHR3.0 is greatly diminished during seconds-long stimulations (especially when using the yellow laser as is done in this work (see Zhang, Chuanqiang, et al. "Optimized photo-stimulation of halorhodopsin for long-term neuronal inhibition." BMC biology 17.1 (2019): 1-17), we remain skeptical of the extent of inhibition during the long manipulations. In short, without a demonstration of effective inhibition throughout the putative consolidation phase (for example by showing a significant decrease in CS frequency throughout the irradiation period), the main claim of the manuscript of phase-specific involvement of CF activity in OKR learning can not be considered to be based on evidence.

Second, the choice of viral targeting strategy leaves gaps in the argument for CF-specific mechanisms. CaMKII promoters are not selective for the IO neurons, and even the most precise viral injections always lead to the transfection of neurons in the surrounding brainstem, many of which project to the cerebellar cortex in the form of mossy fibers (MF). Figure 1Bii shows sparsely-labelled CFs in the flocculus, but possibly also MFs. While obtaining homogenous and strong labeling in all floccular CFs might be impossible, at the very least the authors should demonstrate that their optogenetic manipulation does not affect simple spiking in PNs.

Finally, while the paper explicitly focuses on the effects of CF-evoked complex spikes in the PNs and not, for example, on those mediated by molecular layer interneurons or via direct interaction of the CF with vestibular nuclear neurons, it would be best if these other dimensions of CF involvement in cerebellar learning were candidly discussed.

Reviewer #2 (Public Review):

Summary:

The authors aimed to explore the role of climbing fibers (CFs) in cerebellar learning, with a focus on optokinetic reflex (OKR) adaptation. Their goal was to understand how CF activity influences memory acquisition, memory consolidation, and memory retrieval by optogenetically suppressing CF inputs at various stages of the learning process.

Strengths:

The study addresses a significant question in the cerebellar field by focusing on the specific role of CFs in adaptive learning. The authors use optogenetic tools to manipulate CF activity. This provides a direct method to test the causal relationship between CF activity and learning outcomes.

Weaknesses:

Despite shedding light on the potential role of CFs in cerebellar learning, the study is hampered by significant methodological issues that question the validity of its conclusions. The absence of detailed evidence on the effectiveness of CF suppression and concerns over tissue damage from optogenetic stimulation weakens the argument that CFs are not essential for memory consolidation. These challenges make it difficult to confirm whether the study's objectives were fully met or if the findings conclusively support the authors' claims. The research commendably attempts to unravel the temporal involvement of CFs in learning but also underscores the difficulties in pinpointing specific neural mechanisms that underlie the phases of learning. Addressing these methodological issues, investigating other signals that might instruct consolidation, and understanding CFs' broader impact on various learning behaviors are crucial steps for future studies.

We appreciate the editors and reviewers for their constructive feedback and careful consideration of our manuscript. Despite their acknowledgment of the potential of our study to yield valuable insights into the role of CF activity in cerebellar learning and its phase-specific involvement, we have meticulously addressed all the methodological concerns raised by providing additional clarifications and explanations in this letter.

In response to concerns regarding the efficacy of long-term optogenetic inhibition, we conducted additional in vivo monitoring of CF activity during the irradiation period, confirming sustained inhibition of complex spikes throughout the consolidation phase (Figure 2, lines 112-139). Although stable single-unit recording beyond 40 minutes was not feasible due to technical challenges, the robust suppression of CF-evoked complex spikes we observed during this period (Figure 2, lines 112–139) provides strong evidence that halorhodopsin-mediated inhibition persists over the longer irradiation intervals employed in our behavioral assays.

Moreover, given that there is a concern regarding the CaMKII promoter also inducing expression in neighboring mossy fibers, potentially affecting simple spike activity, we have presented data in Figure 2C, which illustrates that PC simple spike firing rates remain unchanged during prolonged illumination. This finding confirms that our optogenetic manipulation selectively disrupts CF-mediated complex spikes without influencing mossy fiber to PC transmission. We have elucidated these results further in lines 128 to 136.

Lastly, we have broadened our Discussion to consider alternative mechanisms of CF involvement in cerebellar learning, including the modulation of molecular layer interneurons (Rowan et al., 2018) and direct CF interactions with vestibular nuclear neurons (Balaban et al., 1981), thereby offering a more comprehensive perspective on the multifaceted role of CF signaling. Specific clarifications regarding these points are articulated from lines 222 to 242 and 243 to 254 in the manuscript. We are confident that these revisions adequately address the reviewers' concerns and further substantiate the specificity and significance of our study findings

(1) Rowan, Matthew JM, et al. "Graded control of climbing-fiber-mediated plasticity and learning by inhibition in the cerebellum." Neuron 99.5 (2018): 999-1015.

(2) Balaban, Carey D., Yasuo Kawaguchi, and Eiju Watanabe. "Evidence of a collateralized climbing fiber projection from the inferior olive to the flocculus and vestibular nuclei in rabbits." Neuroscience letters 22.1 (1981): 23-29.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Climbing fiber (CF) transmission from the inferior olivary triggers complex spikes (Cs) in Purkinje cells (PCs) driven by a burst of calcium spikes. In the context of motor learning, especially the compensatory optic response, CF transmission serves as instructive signals selectively conveyed to PCs. While the significance of CF input in motor memory formation is widely acknowledged, a comprehensive understanding of its distinct contribution across different temporal windows, spanning from the initial learning phase to the retrieval period, remains incomplete. Therefore, we aimed to investigate the necessity of CF-induced instructive signals in motor learning by assessing their roles in memory acquisition, consolidation, and retrieval. We employed optogenetics to selectively inhibit CF transmission during targeted time windows. Consequently, the absence of CF-induced instructive signals during motor learning impairs memory acquisition. However, when these signals were suppressed during the consolidation and retrieval period, there was neither a loss of long-term memory nor prevention of memory retrieval. Our results highlight that CF transmission plays a specialized and critical role primarily in memory acquisition, rather than subsequent processes.

Revised version: Climbing fiber (CF) transmission from the inferior olive (IO) triggers complex spikes (Cs) in Purkinje cells (PCs) driven by a burst of calcium spikes. In the context of motor learning, especially the compensatory optic response, CF transmission serves as an instructive signal selectively conveyed to PCs. While the significance of CF input in motor memory formation is widely acknowledged, a comprehensive understanding of its distinct contribution across different temporal windows, spanning from the initial learning phase to the retrieval period, remains incomplete. Therefore, we aimed to investigate the necessity of CF-induced instructive signals in motor learning by assessing their roles in memory acquisition, consolidation, and retrieval. We employed optogenetics to selectively inhibit CF transmission during targeted time windows. Consequently, the absence of CF-induced instructive signals during motor learning impairs memory acquisition. However, when these signals were suppressed during the consolidation and retrieval period, there was neither a loss of long-term memory nor a prevention of memory retrieval. Our results highlight that CF transmission plays a specialized and critical role primarily in memory acquisition, rather than in subsequent processes.

Introduction

Earlier version: The cerebellar architecture, characterized by multiple synapses connecting diverse neuron types, enables precise adjustments within intricate neural networks (Boyden et al., 2004). Within this framework, the climbing fiber (CF) transmission derived from the inferior olive (IO) projecting to the cerebellar cortex is generally recognized as the primary pathway for relaying instructive signals to correct errors and enhance motor performance, including oculomotor behaviors (Albus, 1971; Frens et al., 2001; Ito, 1972; Marr, 1969).

Revised version: The cerebellar architecture, characterized by multiple synapses connecting diverse neuron types, enables precise adjustments within intricate neural networks (Boqhyden et al., 2004). Within this framework, the climbing fiber (CF) transmission derived from the inferior olive (IO) projecting to the cerebellar cortex is generally recognized as the primary pathway for relaying instructive signals to correct errors and enhance motor performance, including oculomotor behaviors (Albus, 1971; Frens et al., 2001; Ito, 1972; Marr, 1969).

A prominent example of cerebellum-dependent motor learning is the optokinetic reflex (OKR), a compensatory optic response evoked by an unstable visual field (Cahill & Nathans, 2008; Schweigart et al., 1997). In OKR, adaptation to stabilize retinal images amid mismatched visual input requires the engagement of multiple neural substrates. For instance, complex spikes (Cs) in floccular Purkinje cells (PCs) driven by CF input, emerge in response to unpredicted motor performance errors (Clopath et al., 2014; Eccles et al., 1966; Goossens et al., 2004; Palmer et al., 2010). Rotational visual stimulation induces Cs discharge in floccular PCs, reflecting CF-induced Cs generation in response to error signals generated from errors in oculomotor performance (Frens et al., 2001; Graf et al., 1988). Therefore, this fine-tuning of Cs in PCs is expected to regulate learning by modulating error-correcting adaptations (Yang & Lisberger, 2014).

Previous studies highlight a potential duality in CF involvement in cerebellar motor learning. For example, one study using a modified behavioral paradigm to suppress Cs suggested that simple spikes (Ss) alone may suffice for motor learning (Ke et al., 2009). In contrast, pharmacological lesioning and optogenetic inactivation of the IO impair oculomotor learning (Pham et al., 2020) and prevent eye blink conditioning (Kim et al., 2020; Medina et al., 2002; Welsh & Harvey, 1998). While CF’s role in acquisition is well-supported by existing literature, its involvement in downstream memory processes like consolidation and retrieval remains contentious. This ambiguity stems from methodological limitations in previous studies, which have lacked the spatial and temporal precision to isolate CF function at distinct learning stages. Addressing this knowledge gap is essential to clarify whether CF input is uniformly critical or phase-specific in motor learning and memory.

To address this gap, we investigated the necessity of CF-induced instructive signals in motor learning by assessing their contributions across memory acquisition, consolidation, and retrieval phases. Using optogenetics, we selectively inhibited CF signals in the flocculus, a lobe-like cerebellar structure primarily responsible for ocular reflex, during targeted learning stages. Our findings provide compelling evidence for a phase-specific role of CF transmission in motor learning. This indicates that CF input is essential primarily during acquisition, with minimal involvement in consolidation and retrieval.

Results

Optogenetic inhibition of climbing fiber transmission derived from the inferior olive

To specifically manipulate CF signals in the cerebellar cortex, generally known to originate from the IO, we injected AAV1-CaMKIIα-eNpHR 3.0-EYFP or AAV1-CaMKIIα-EGFP into the IO of wild-type mice (Figure 1A, B). This approach allowed visualization of IO neuron expression and CF terminals in the brainstem nucleus and flocculus. Importantly, the selectively segregated expression in these regions enabled manipulation of CF terminals without disrupting IO somas, which, if damaged, can critically impair motor performance and ocular reflexes due to the IO’s role as a prominent center of sensory integration (Pham et al., 2020; Shaikh et al., 2017; Van Der Giessen et al., 2008).

Optogenetic inhibition of climbing fiber transmission derived from the inferior olive.

A. Schematic diagram showing virus injection for optogenetic manipulation of CF transmission.

B. Bi. A virus is expressed specifically only in the IO region (IO, inferior olive; py, pyramidal tract; ml, medial lemniscus). Scale bar, 100 μm. Bii. Virus expression along CF in the cerebellar flocculus region. (GCL, granule cell layer; PCL, Purkinje cell layer; ML, molecular layer; FL, flocculus). Scale bar, 100 μm.

C. A scheme of optogenetic suppression using a yellow laser (593 nm, 3–5 mW) while recording CF EPSCs under voltage clamp mode.

D. Left, representative traces of CF EPSCs of the GFP (top) and NpHR group (bottom) with (dark line) and without (light line) opto-stimulation. Right, Quantitative analysis of suppression of CF EPSCs (GFP: n = 10 cells/2 mice, NpHR: n = 9 cells/5 mice, *** p < 0.001, Unpaired t-test).

E. A scheme of optogenetic suppression while recording complex spikes under current clamp mode.

F. Representative traces of complex spike of the GFP (top) and NpHR group (bottom) with (dark line) and without (light line) opto-stimulation. Quantitative analysis of number of complex spike spikelets (GFP: n = 8 cells/1 mouse, NpHR: n = 8 cells/5 mice, ** p = 0.002; Paired t-test).

To confirm the effectiveness of optogenetic inhibition of climbing fiber (CF) transmission, we performed whole-cell patch-clamp recordings of Purkinje cells (PCs) in acute slices of cerebellar vermis lobules 4–5. A 593 nm yellow laser was used to activate halorhodopsin (NpHR). PCs were selected based on optical fluorescence expression (Figures 1C and 1E), and a stimulation electrode was positioned in the granule cell layer near the recorded PC. The stimulation intensity was carefully calibrated to evoke minimal CF excitatory postsynaptic currents (EPSCs). Laser illumination (3–5 mW) during stimulation resulted in a significant reduction in EPSC amplitude, indicating robust inhibition of CF transmission. Notably, this inhibition was consistent, with 9 out of 9 cells in the NpHR group exhibiting reliable suppression without failures. Importantly, laser illumination did not affect CF EPSCs in the GFP group (Figure 1D). Under the current-clamp mode, CF stimulation-induced spikelets were also completely abolished in the NpHR group but remained unaffected in the GFP group (Figure 1F). These findings establish that optogenetic inhibition effectively suppresses CF-PC synaptic transmission in vitro.

To extend our findings in vivo, we recorded from PCs in awake, head-fixed mice during 593 nm optogenetic stimulation. AAV1-CaMKIIα-eNpHR 3.0-EYFP or AAV1-CaMKIIα- EGFP was injected into the IO followed by the creation of a cranial window above the cerebellar vermis three weeks later (Figure 2A). Due to technical constraints preventing direct access to the flocculus for in vivo recordings, we targeted cerebellar lobule IV/V. PCs were identified based on the detection of Cs (Figure 2B).

Earlier version: Optogenetic inhibition of climbing fiber transmission bl ked Cs firing but not Ss firing in vivo

Revised version: Optogenetic inhibition of climbing fiber transmission blocked Cs firing but not Ss firing in vivo

Earlier version: A. Schematic diagram of virus injection into IO and in vivo recordin ring optogenetic manipulation of CF transmission. Arrow heads indicate NpHR-expressing CF terminals in the molecular layer of the cerebellar cortex. Scale bar, 50 μm

Revised version: A. Schematic diagram of virus injection into IO and in vivo recording during optogenetic manipulation of CF transmission. Arrow heads indicate NpHR-expressing CF terminals in the molecular layer of the cerebellar cortex. Scale bar, 50 μm

B. Representative electrophysiological recording traces of PC firing from the GFP and NpHR groups, without (OFF) and with (ON) yellow (593 nm) opto-stimulation. The upper plots show 10-second traces without and with opto-stimulation, and the lower plots display a 1-second segment from the respective upper traces. Asterisks denote Cs

C. Comparison of Ss firing rates during optogenetic inhibition. For the GFP group, firing rates with laser off (n = 11 cells from 4 mice) versus on (n = 11 cells from 4 mice) were not significantly different (Mann-Whitney test, P = 0.8594). Similarly, for the NpHR group, no significant difference was observed between OFF (n = 10 cells from 6 mice) and ON (n = 10 cells from 6 mice) conditions (Mann-Whitney test, p > 0.9999).

D. Comparison of Cs firing rates during optogenetic inhibition. In the GFP group, no significant difference was observed between OFF (n = 8 cells from 4 mice) and ON (n = 8 cells from 4 mice) conditions (Mann-Whitney test, p = 0.4563). In the NpHR group, Cs firing rates were significantly reduced during opto-stimulation (OFF, n = 6 cells from 6 mice; ON, n = 6 cells from 6 mice; Mann-Whitney test, ** p = 0.0022). Of note, 6 out of 8 recorded cells were responsive to optogenetic suppression.

Optogenetic inhibition selectively suppressed Cs firing rates in the NpHR group without affecting Ss rates, whereas no changes were observed in the GFP control group (Figures 2C and 2D). Recordings from the same neurons during both laser-off and laser-on conditions demonstrated consistent and robust suppression of Cs activity throughout the 40-minute optostimulation period. Although continuous recordings over several hours were not feasible, the stability and sustained suppression observed at 40 minutes strongly suggest that the manipulation would remain effective during the extended durations required for behavioral experiments. These findings confirm that optogenetic inhibition of NpHR-expressing CFs selectively suppresses Cs activity in awake mice without altering Ss activity. Furthermore, they validate the feasibility and specificity of using optogenetic interventions to block CF transmission both in vitro and in vivo, reinforcing the reliability of this approach for studying CF contributions to cerebellar function and behavior.

Inhibition of climbing fiber transmission during memory acquisition

An evaluation of CF’s direct contribution in a region-specific manner has not been performed. Thus, after confirming the viability of optogenetic inhibition of CF, we investigated whether the elimination of CF transmission in the cerebellar flocculus affects cerebellar motor learning. When images on the retina are unstable due to mismatched visual inputs, CF signals are conventionally perceived to be responsible for sending error signals to the cerebellar cortex to improve ocular reflexes (Bloedel & Bracha, 1998; Frens et al., 2001; Zang & De Schutter, 2019). Using the same mouse model shown in Figure 1, we surgically implanted an optical fiber in flocculus to suppress CF transmission in vivo during OKR (Figure 3A). Head- fixed mice were placed in a rotating drum painted with a black striped pattern to provide visual input, and its rotating motion elicited an unstable retinal image. Hence, in the beginning, the trace of the mouse’s eye movement represented in a sine curve was only partially overlapping with the trace of the drum’s kinetics, but as the mouse underwent more sessions of OKR learning, the trace of eye movement became almost identical (Figure 3Bi). Notably, the group with inhibited CF transmission exhibited only slight changes or no improved trace of eye movements (Figure 3Bii). Overall, the results demonstrated that CF transmission is required for motor memory acquisition.

Inhibition of climbing fiber transmission during memory acquisition of OKR.

Earlier version: A. Illustration of the optokinetic reflex behavioral test (left) an experimental scheme (right). For training, visual stimulation was given by constant rotat n of the screen, while the table, where a mouse’s head is fixed, remained stationary. r day 1, a mouse was trained for a total of 50 minutes, and on day 2, an amount of sustained memory was validated via gain check. Scale bar, 100 μm

Revised version: A. Illustration of the optokinetic reflex behavioral test (left) and experimental scheme (right). For training, visual stimulation was given by constant rotation of the screen, while the table, where a mouse’s head is fixed, remained stationary. For day 1, a mouse was trained for a total of 50 minutes, and on day 2, an amount of sustained memory was validated via gain check. Scale bar, 100 μm

B. Representative traces of the screen (black sinusoidal curve) and eye movements (grey sinusoidal curve) divided into the GFP (i) and NpHR (ii) groups. Top traces (pre) were acquired before learning, and bottom traces (post) were obtained after 50 min of learning. The vertical scale bar represents 10 degrees per second and the horizontal scale bar represents 0.5 seconds.

C. The learning curve from 0 min to 50 min, and +24 hr period. Change of gain is indicated at each time point (+10 min increment). Yellow boxes indicate opto-stimulation (12.5 mW for 10 min/session, total of 5 sessions). The left graph represents the GFP (n = 8) and the right graph indicates the NpHR (n = 9) group.

D. Comparison of gain changes between the GFP and NpHR groups (GFP, n = 8 mice; NpHR, n = 9 mice). Percentages of gain increment from 0 min to 50 min were calculated. The change of gain of the GFP group was significantly larger than that of the NpHR group (left; Unpaired t-test, p = 0.0007). The gain of the GFP group was significantly increased from 0 min to 50 min (right; Paired t-test, **** p < 0.0001), while the NpHR group showed no significant improvement of gain from 0 min to 50 min (right; Paired t- test, p = 0.1046).

This indicates that CF transmission is essential for cerebellar motor learning by presenting two contrasting learning curves. The control group exhibited an increasing trend in the learning curve, whereas the NpHR group’s learning curve was almost flat, which reflects no change of gain (Figure 3C). The gain value through the learning sessions was distinctly increased in the control group, whereas the NpHR group showed no significant change in gain (Figure 3D). Furthermore, from 0 to 50 min period of learning, the control group demonstrated almost a doubling in gain, whereas the NpHR group showed little or almost no change of gain (Figure 3E). These results indicate that CF transmission is required for the proper acquisition of motor memory.

Inhibition of climbing fiber transmission during memory consolidation and retrieval phase

When memories are initially acquired, they are generally unstable; hence, it requires some time to shape a stabilized long-term memory (Attwell et al., 2002). Therefore, memories undergo multiple processes that allow distinct shifts from short-term to long-term memory, leading to the construction of adapted behavior (Shutoh et al., 2006). Having confirmed the active involvement of CF transmission during memory acquisition, we sought to understand if its role is limited to the acquisition period or extends to memory transfer during the consolidation phase for long-term storage. We implemented identical surgical and behavioral procedures as previously conducted. However, instead of optogenetically inhibiting the transmission during memory acquisition, we suppressed CF activity in two distinct optogenetic manipulation schemes: one targeting short-term inhibition (30 min) and the other for long-term inhibition (6 hr) (Figures 4A and 4D). This approach was based on findings that long-term memories associated with cerebellum-dependent behaviors are vulnerable to disruption within specific time windows of consolidation, spanning from minutes to hours (Cooke et al., 2004; Titley et al., 2007). Learning curves confirmed the successful gain enhancement, validating the integrity of memory acquisition (Figures 4B and 4E). Once the gain was robustly enhanced, optogenetic inhibition of CF transmission during consolidation did not yield any change in the maintenance of memory, irrespective of the duration of manipulation. Consequently, long-term memory was well-established in both NpHR and GFP groups (Figures 4C and 4F). In addition, we verified whether CF signaling is required during the retrieval phase (Figure 4G). When mice, fully trained in the optokinetic reflex, displayed increased gain, disrupting CF signaling did not affect the retrieval of motor memory (Figure 4H). This divergence underscores a specialized role for CF transmission in the early stages of motor learning, with subsequent phases relying on distinct neural mechanisms. The results highlight phase-specific contributions of CF-driven error signaling, aligning with the hypothesis that memory consolidation and retrieval depend on downstream circuits independent of CF activity.

Inhibition of climbing fiber transmission during memory consolidation and retrieval phase.

A. Illustration of the experimental scheme. 10 min of optokinetic reflex training was repeated 5 times, and then opto-stimulation (12.5 mW) was given at 0 min post-learning for 30 min.

B. The learning curve from 0 min to 50 min, and +24 hr period. Change of gain is indicated at each time point (+10 min increment). A yellow box indicates opto-stimulation at 0 min post-learning. The left graph represents the GFP and the right graph indicates the NpHR group.

C. Comparison of consolidation percentage between the GFP and NpHR groups (GFP group, n = 7 mice; NpHR group, n = 9 mice). The percentage of sustained memory after 24 hr was calculated. No difference was observed between the groups (p = 0.3173; Unpaired t- test).

D. Illustration of the experimental scheme. 10 min of optokinetic reflex training was repeated 5 times, and then opto-stimulation (12.5 mW) was given at 0 min post-learning for 6 hr.

E. The learning curve from 0 min to 50 min, and +24 hr period. Change of gain is indicated at each time point (+10 min increment). A yellow box indicates opto-stimulation at 0 min post-learning. The left graph represents the GFP group, and the right graph indicates the NpHR group.

F. Comparison of gain changes between GFP and NpHR groups (GFP group, n = 6 mice; NpHR group, n = 6 mice). Percentages of gain increment from 0 min to 50 min were calculated. No difference was observed between the groups (Unpaired t-test, p = 0.6405).

G. Illustration of the experimental scheme. Following proper motor learning on day 1, the remaining gain value was checked on day 2, after testing 1st gain retrieval without opto-stimulation, 2nd gain retrieval was tested with opto-stimulation to validate the effect of CF inhibition on gain retrieval (opto-stimulation was given 3 times for 24 seconds each at 12.5 mW).

H. Gain retrievals before and during opto-stimulation were compared. The left graph represents the GFP group, and the right shows the NpHR group (GFP group, n = 6 mice; NpHR group, n = 8 mice). There was no significant variation between the groups (NpHR group, p = 0.2439; GFP group, p = 0.3180; Paired t-test).

Discussion

Motor learning is a finely tuned process, yet its underlying cellular mechanisms operate with remarkable precision, adapting selectively to task demands. This adaptability relies on CF activity, originating from the inferior olive, which instructs cerebellar circuits to reduce motor errors and enhance performance. CF-driven Cs in PCs have long been associated with motor learning, but how CF signals vary across learning stages remains less understood. Our study uncovers a distinct role for CF signaling that challenges traditional views: CF input is essential for driving motor adaptation during memory acquisition but surprisingly dispensable during consolidation and retrieval. This phase-specific pattern of CF transmission in cerebellum-dependent learning indicates that its role is not uniform over time, as previously thought, but instead selectively tunes cerebellar circuits to optimize adaptation.

Earlier version: Through optogenetic inhibition of CF transmission, we found that CF activity is critical for enhancing motor gain during initial memory acquisition, where substantial adjustments and error corrections are required. In contrast, CF input played a minimal role during memory consolidation and retrieval, suggesting that other mechanisms may take precedence as motor memories stabilize and transfer to downstream circuits. This discovery underscores a refined model of CF function: CFs initiate learning adjustments but are not integral to later stages of memory storage or retrieval, reshaping our understanding of the cerebellum’s contribution to motor memory over time.

Revised version: By inhibiting CF transmission optogenetically, we demonstrated that CF activity is crucial for enhancing motor gain in the early phases of memory acquisition, where substantial adjustments and corrections are necessary. Conversely, CF input was found to have a minor influence during memory consolidation and retrieval, indicating that different mechanisms may dominate as motor memories become stable and transition to downstream circuits. This finding refines our understanding of CF function: CFs prompt learning adjustments but do not play a key role in the later phases of memory storage or recall, thereby altering our perspective on the cerebellum’s role in memory formation over time.

Earlier version: Motor memories acquired through cerebellar-dependent learning are believed to transfer to the vestibular nuclei for long-term storage (Jang et al., 2020; Seo et al., 2024; Shutoh et al., 2006). By applying optogenetic inhibition of CF transmission at various learning stages, we differentiated the phase-specific role of CF signaling. CF activity proved indispensable for motor adaptation (i.e., gain increase) during acquisition, yet was unnecessary during consolidation and retrieval, suggesting CF input’s primary role is in initial memory formation rather than in later memory processing.

Revised version: Motor memories gained through cerebellar-dependent learning are thought to transfer to the vestibular nuclei for long-term retention (Jang et al., 2020; Seo et al., 2024; Shutoh et al., 2006). By employing optogenetic inhibition of CF transmission at various stages of learning, we identified the phase-specific role of CF signaling. CF activity was crucial for motor adaptation, such as gain increase training, during the acquisition phase but was not needed during consolidation and retrieval. This suggests that the primary function of CF input lies in initial memory formation rather than in subsequent memory processing.

Earlier version: Our findings also indicate that CF-dependent and CF-independent mechanisms work in a context-sensitive manner during cerebellar learning, consistent with established insights from the vestibulo-ocular reflex (VOR) and OKR systems. Specifically, CF signaling appears essential in tasks involving substantial, persistent errors that require precise error correction. For instance, in the VOR, gain increases depend on CF-driven complex spikes, which provide adaptive signals in response to large error inputs (Ke et al., 2009). Conversely, CF-independent processes can support VOR gain decreases, where smaller or more gradual adjustments are sufficient. This selective role of CF signaling reflects a broader functional adaptation within cerebellar circuits, allowing different pathways to be recruited according to motor learning demands. Our OKR findings align with this, highlighting CF-dependent pathways in tasks with significant error correction during acquisition, while CF-independent mechanisms likely support incremental, lower-error adjustments, facilitating efficient motor adaptation without high-frequency CF input. This division of labor between CF-dependent and CF-independent processes underscores the cerebellum’s adaptive plasticity, where CF input selectively engages based on the scale and persistence of motor errors. By dynamically alternating between these mechanisms, the cerebellum achieves a balance between precision and flexibility, optimizing motor adaptation across various behavioral contexts. This interplay provides a versatile framework for motor learning across cerebellar-dependent tasks, extending from VOR to OKR and likely other sensorimotor adaptations.

Revised version: Our results suggest that CF-dependent and CF-independent mechanisms may operate in a context-sensitive manner during cerebellar learning, consistent with established insights derived from the vestibulo-ocular reflex (VOR) and OKR systems. Specifically, CF signaling appears essential in tasks involving substantial, persistent errors requiring precise error correction. For instance, in the VOR, the increase in gain is contingent upon CF-driven complex spikes, which provide adaptive signals in response to large error inputs (Ke et al., 2009). Conversely, CF-independent processes can facilitate reductions in VOR gain, wherein smaller or more gradual adjustments are adequate. This selective role of CF signaling reflects a broader functional adaptation within cerebellar circuits, permitting the engagement of different pathways according to the demands of motor learning. The findings related to our OKR results correspond with this notion, emphasizing the involvement of CF-dependent pathways in tasks characterized by substantial error correction throughout the acquisition phase. In contrast, CF-independent mechanisms likely enable incremental adjustments of lesser magnitude, thereby promoting efficient motor adaptation in the absence of CF input. This delineation of responsibilities between CF-dependent and CF-independent processes highlights the cerebellum’s adaptive plasticity, wherein CF input is selectively activated based on the magnitude and persistence of motor errors. The cerebellum maintains a balance between precision and flexibility by dynamically alternating between these mechanisms, enhancing motor adaptation across various behavioral contexts. This interplay offers a versatile framework for motor learning in cerebellar-dependent tasks, encompassing VOR, OKR, and likely other forms of sensorimotor adaptation.

Added in the revised version: In addition to their robust excitatory drive on PCs, CFs shape cerebellar computation by recruiting molecular-layer interneurons to impose feedforward inhibition. Optogenetic activation of molecular-layer interneurons, which silences PCs, markedly suppresses CF- evoked dendritic calcium spikes, thereby CF transmission may dynamically gate their own instructive signals, shifting the polarity and magnitude of synaptic plasticity at PF–PC synapses (Rowan et al., 2018). Beyond the cerebellar cortex, anatomical tracing reveals that the same IO neurons that innervate the flocculus also send axon collaterals to the vestibular nuclei (Balaban et al., 1981), suggesting that CF signals can simultaneously modulate cortical and vestibular outputs. While the behavioral consequences of this direct CF-vestibular nucleus pathway remain to be defined, it provides a parallel channel through which IO teaching signals may coordinate cortical plasticity with adaptive vestibular reflexes.

Earlier version: The process of systems consolidation, facilitated by the medial vestibular nuclei (MVN) and downstream circuitry of floccular PCs, is essential for long-term memory storage (Shutoh et al., 2006). Given that CF activity modulates floccular PC firing, we hypothesized that CF signaling might indirectly influence MVN activity during memory consolidation and retrieval. However, contrary to this expectation, our results revealed no substantial effect of CF transmission on long-term memory formation or recall, suggesting that CF signaling does not directly impact MVN function. Instead, previous studies suggest that PC intrinsic plasticity supports memory consolidation (Seo et al., 2024), while granule cell activity contributes to memory retrieval (Wada et al., 2014). These findings reinforce the notion that CF signaling primarily facilitates the acquisition phase of motor memory, while simple spiking in PCs may encode timing and frequency information crucial for stabilizing memory traces in downstream circuits. Ss could either complement or function independently of CF transmission, especially during phases when error correction demands are minimal. Mechanistically, suppressing CF signaling likely mimicked key aspects of long-term depression (LTD) at CF-PC synapses, such as reducing the amplitude of Cs (Hansel & Linden, 2000). Under normal conditions, CF input promotes parallel fiber-PC (PF-PC) LTD by driving calcium influx during error correction, such as in response to retinal slip (Coesmans et al., 2004). CF-LTD modifies Cs waveforms and attenuates CF-evoked dendritic calcium transients (Weber et al., 2003). In our study, optogenetic CF suppression likely simulated the functional consequences of CF-LTD, which may have impaired PF-PC LTD and subsequently hindered gain enhancement during learning (Han et al., 2007).

Revised version: The process of systems consolidation, facilitated by the medial vestibular nuclei (MVN), a downstream circuitry of floccular PCs, is essential for long-term memory storage (Shutoh et al., 2006). Given that CF activity modulates floccular PC firing, we hypothesized that CF signaling might indirectly influence MVN activity during memory consolidation and retrieval. However, contrary to this expectation, our results revealed no substantial effect of CF transmission on long-term memory formation or recall, suggesting that CF signaling does not directly impact MVN function. Instead, previous studies suggest that PC intrinsic plasticity supports memory consolidation (Seo et al., 2024), while granule cell activity contributes to memory retrieval (Wada et al., 2014). These findings reinforce the notion that CF signaling primarily facilitates the acquisition phase of motor memory, while simple spiking in PCs may encode timing and frequency information crucial for stabilizing memory traces in downstream circuits. Ss could either complement or function independently of CF transmission, especially during phases when error correction demands are minimal. Mechanistically, suppressing CF signaling likely mimicked key aspects of long-term depression (LTD) at CF-PC synapses, such as reducing the amplitude of Cs (Hansel & Linden, 2000). Under normal conditions, CF input promotes parallel fiber-PC (PF-PC) LTD by driving calcium influx during error correction, such as in response to retinal slip (Coesmans et al., 2004). CF-LTD modifies Cs waveforms and attenuates CF-evoked dendritic calcium transients (Weber et al., 2003). In our study, optogenetic CF suppression likely simulated the functional consequences of CF-LTD, which may have impaired PF-PC LTD and subsequently hindered gain enhancement during learning (Han et al., 2007).

Added in the revised version: Our findings underscore a phase-specific role for CF input, demonstrating that CF activity is critical for motor memory acquisition but less necessary for consolidation and retrieval.

Earlier version: Our findings underscore a phase-specific role for climbing fiber (CF) input, demonstrating that CF activity is critical for motor memory acquisition but less necessary for consolidation and retrieval. Using optogenetic inhibition, we achieved precise temporal control over CF activity within distinct learning phases, enabling targeted manipulation of CF signaling. This differs from prior studies that employed pharmacological or lesion methods, which affect broader cerebellar pathways and may implicate CF signaling across multiple memory stages. The task demands inherent to different learning phases may explain these observed differences. Tasks requiring high-error correction, such as those encountered during the active learning phase, likely depend on CF-driven signals for precise error correction during acquisition. In contrast, CF-independent mechanisms may dominate later phases where such adjustments are less critical. This context-sensitive role highlights the adaptability of CF signaling in tuning cerebellar circuits to meet the specific demands of each learning phase. Moreover, cerebellar-dependent memory consolidation exhibits biphasic temporal dynamics: an initial rapid phase of consolidation followed by prolonged stabilization. Our study specifically focused on the early consolidation phase, emphasizing cerebellar plasticity mechanisms that are critical for encoding and stabilizing motor memory (Cooke et al., 2004; Seo et al., 2024; Steinmetz & Freeman, 2016; Titley et al., 2007). Importantly, changes during this early period significantly shape memory formation and retention, while alterations beyond this timeframe have minimal impact on storage or performance. While our findings align with this framework, extended consolidation periods may involve slower processes, such as systems-level integration into downstream nuclei or cortical regions, contributing to memory stabilization or generalization. Future investigations using chronic in-vivo recordings or live imaging could elucidate these prolonged dynamics and clarify CF transmission’s nuanced contributions across phases of cerebellar-dependent memory consolidation.

Revised version: Using optogenetic inhibition, we achieved precise temporal control over CF activity within distinct learning phases, enabling targeted manipulation of CF signaling. This differs from prior studies that employed pharmacological or lesion methods, which affect broader cerebellar pathways and may implicate CF signaling across multiple memory stages. The task demands inherent to different learning phases may explain these observed differences. Tasks requiring high-error correction, such as those encountered during the active learning phase, likely depend on CF-driven signals for precise error correction during acquisition. In contrast, CF-independent mechanisms may dominate later phases where such adjustments are less critical. This context-sensitive role highlights the adaptability of CF signaling in tuning cerebellar circuits to meet the specific demands of each learning phase. Moreover, cerebellar- dependent memory consolidation exhibits biphasic temporal dynamics: an initial rapid phase of consolidation followed by prolonged stabilization. Our study specifically focused on the early consolidation phase, emphasizing cerebellar plasticity mechanisms that are critical for encoding and stabilizing motor memory (Cooke et al., 2004; Seo et al., 2024; Steinmetz & Freeman, 2016; Titley et al., 2007). Importantly, changes during this early period significantly shape memory formation and retention, while alterations beyond this timeframe have minimal impact on storage or performance. While our findings align with this framework, extended consolidation periods may involve slower processes, such as systems-level integration into downstream nuclei or cortical regions, contributing to memory stabilization or generalization. Future investigations using chronic in-vivo recordings or live imaging could elucidate these prolonged dynamics and clarify CF transmission’s nuanced contributions across phases of cerebellar-dependent memory consolidation.

In conclusion, inhibiting CF transmission during motor learning impairs gain increases specifically during acquisition. These findings provide critical insights into CF signaling’s role in cerebellum-dependent learning while underscoring the need for further exploration of how cerebellar circuits contribute to adaptive motor memory over extended timeframes.

Key resources table

Animals and Stereotaxic Surgery for Virus Injection

All experimental procedures were approved by the Seoul National University Institutional Animal Care and Use Committee. First, wild-type mice aged 7–10 weeks were anesthetized using intraperitoneal injections of a Zoletil/Rompun mixture (30 mg / 10 mg/kg). For optogenetic expression in CF, the virus was injected into the IO 2–3 weeks before the behavioral test, as previously described (Kimpo et al., 2014; Roh et al., 2020). Briefly, bilateral injections were made at the midpoint between the edge of the occipital bone and the C1 cervical vertebra. The glass pipette was set at a 55° angle from the vertical and 7° from the midline. After approaching a 3 mm depth, a virus solution containing 100–200 nL of AAV1-CaMkIIα-eNpHR3.0-EYFP or AAV9.CaMkIIα-EGFP was injected with a Picopump at 5 nL/s. The pipettes were left in place for 10 min before they were removed to minimize backflow.

Stereotaxic Surgery for Optic Cannula Implantation

After the mice were fully recovered (a few days to a week) from the first surgery of virus injection, a second operation of optic cannula implantation (Ø1.25 mm Ceramic Ferrule, L=3.0 mm) near cerebellar flocculus and head fixation was conducted. Unlike the fiber optic cannula, head fixation parts were hand-made. Before the surgery, mice were anesthetized by intraperitoneal injection of Zoletil 50 (Virbac, 15 mg/kg) and xylazine (Rompun, Bayer, 15 mg/kg) mixture. Mice were given 48 hours to recover after surgery.

Confirmation Viral Expression

After completion of all the experiments, sampling using cardiac perfusion is performed. The brain was then extracted and coronal sections were made at 30 μm intervals. Images were acquired and processed using a confocal microscope (Zeiss LSM 7 MP, Carl Zeiss, Jena, Germany) and Zen software (Zeiss). The location of virus expression was confirmed by comparison with the brain atlas (Paxinos and Franklin’s The Mouse Brain in Stereotaxic Coordinates 4th Edition by George Paxinos).

Optical suppression of CF terminals in the cerebellar cortex in vivo

Earlier version: To manipulate CF transmission-induced Cs in PCs, AAV1.CaMKII.eNpHR3.0.eYFP or AAV8.CaMKIIa.eGFP was bilaterally injected into the IO. A glass micro-pipette was used to inject 200 nL per site. To optically suppress the CF terminals and record CF-PC synaptic response, a small cranial window was made above the cerebellar cortex lobule IV/V a week after the viral injection. After 2 weeks from the viral injection, mice underwent single-unit recordings. The recordings were performed as described in Neural data acquisition and analysis. Cs responses were acquired during 10 seconds of yellow on and off sessions 20 times to test the effect of CF terminal suppression on CF-PC synaptic transmission.

Revised version: To manipulate CF transmission-induced Cs in PCs, AAV1.CaMKIIl7l.eNpHR3.0.eYFP or AAV8.CaMKIIa.eGFP was bilaterally injected into the IO. A glass micro-pipette was used to inject 200 nL per site. To optically suppress the CF terminals and record CF-PC synaptic response, a small cranial window was made above the cerebellar cortex lobule IV/V a week after the viral injection. After 2 weeks from the viral injection, mice underwent single-unit recordings. The recordings were performed as described in Neural data acquisition and analysis. Cs responses were acquired during 10 seconds of yellow on and off sessions 20 times to test the effect of CF terminal suppression on CF-PC synaptic transmission.

Neural data acquisition and analysis

Single-unit recordings were performed while mice were head-fixed and awake in the recording rack. A Digital Lynx system (Neuralynx, Bozeman, MT) was used to amplify, band pass filter (0.1–8000 Hz for simple spikes, 10–200 Hz for complex spikes), and digitize the electrode recordings. A silicon neural probe (Cambridge NeuroTech) was used to probe and acquire Cs signals. Cs signals were collected at a 32 kHz sampling rate. Cs waveforms were manually selected based on visual inspection of the averaged appearance of characteristic negative Ca2+ peaks.

Slice preparation, electrophysiology, and optogenetic validation

An acute brain slice preparation and electrophysiological experiments were carried out as previously described (Lee et al., 2024). First, mice aged 9–10 weeks with AAV1.CaMKIIα.eNpHR 3.0.WPRE-EYFP injection into their inferior olivary neuron 2–3 weeks before the preparation were anesthetized by isoflurane and briefly decapitated. Then, 250 μm-thick sagittal slices of the cerebellar vermis were obtained from the mice using a vibratome (VT1200S, Leica). The ice-cold cutting solution contained 75 mM sucrose, 75 mM NaCl, 2.5 mM KCl, 7 mM MgCl2, 0.5 mM CaCl2, 1.25 mM NaH2PO4, 26 mM NaHCO3, and 25 mM glucose with bubbled 95% O2 and 5% CO2. The slices were immediately transferred into the artificial cerebrospinal fluid (ACSF) containing 125 mM NaCl, 2.5 mM KCl, 1 mM MgCl2, 2 mM CaCl2, 1.25 mM NaH2PO4, 26 mM NaHCO3, and 10 mM glucose with bubbled 95% O2 and 5% CO2. They were allowed to recover at 32 °C for 30 min and at room temperature for 1 hour. All recordings were performed within 8 hr of recovery.

Earlier version: Brain slices were placed in a submerged chamber with perfusion of ACSF for at least 10 min before recording. We used OptoPatcher (HEKA Elektronik) holding recording pipette (2–3 MΩ) filled with internal solution containing 9 mM KCl, 10 mM KOH, 120 mM K-gluconate, 3.48 mM MgCl2, 10 mM HEPES, 4 mM NaCl, 4 mM Na2ATP, 0.4 mM Na3GTP, and 17.5 mM sucrose (pH 7.25). The stimulation electrode was placed on the granule cell layer near the PC. The stimulation isolator injected a brief current pulse to the stimulation electrode and was controlled using PatchMaster software (HEKA Elektronik). Stimulation intensity (6–30 μA) was calibrated to minimally evoke CF EPSC. Regarding the laser on trial, a 593.5 nm irradiation from a yellow laser was applied 5 s before CF stimulation. The eEPSC amplitude was averaged over three lasers on and off trials. Under current clamp mode, complex spikes induced by CF stimulation were recorded using the same validation protocol. Overall, the laser illumination (3-5 mW) consistently inhibited CF transmission. Electrophysiological data were acquired using an EPC10 patch-clamp amplifier (HEKA Elektronik) and PatchMaster software (HEKA Elektronik) with a sampling frequency of 20 kHz, and the signals were filtered at 2 kHz. All electrophysiological recordings were acquired from the central cerebellar vermis. The amplitude of EPSC was analyzed using Igor Pro (WaveMetrics).

Revised version: Brain slices were placed in a submerged chamber with perfusion of ACSF for at least 10 min before recording. We used OptoPatcher (HEKA Elektronik) holding recording pipette (2–3 MΩ) filled with internal solution containing 9 mM KCl, 10 mM KOH, 120 mM K- gluconate, 3.48 mM MgCl2, 10 mM HEPES, 4 mM NaCl, 4 mM Na2ATP, 0.4 mM Na3GTP, and 17.5 mM sucrose (pH 7.25). The stimulation electrode was placed on the granule cell layer near the PC. The stimulation isolator injected a brief current pulse to the stimulation electrode and was controlled using PatchMaster software (HEKA Elektronik). Stimulation intensity (6–30 μA) was calibrated to minimally evoke CF EPSC. Regarding the laser on trial, a 593.5 nm irradiation from a yellow laser was applied 5 s before CF stimulation. The eEPSC amplitude was averaged over three lasers on and off trials. Under current clamp mode, complex spikes induced by CF stimulation were recorded using the same validation protocol. Overall, the laser illumination (3-5 mW) consistently inhibited CF transmission.

Added in the revised version: Electrophysiological data were acquired using an EPC10 patch-clamp amplifier (HEKA Elektronik) and PatchMaster software (HEKA Elektronik) with a sampling frequency of 20 kHz, and the signals were filtered at 2 kHz. All electrophysiological recordings were acquired from the central cerebellar vermis. The amplitude of EPSC was analyzed using Igor Pro (WaveMetrics).

Optokinetic reflex behavior test

First, two sessions of acclimation were conducted. During each session, the mouse was restrained using a custom-made restrainer for 20 min both with and without light for habituation to the recording environment. Following the procedure reported by (Stahl et al., 2000), a calibration was performed. This allowed for the conversion of the dynamics of pupil- to-eye rotation. Next, three basal oculomotor performances including (OKR), vestibulo- ocular reflex in the dark (dVOR), and vestibulo-ocular reflex in the light (lVOR) were measured. During the recording of eye movements, it was necessary to control pupil dilation that was induced by an absence of light in dark conditions. Thus, physostigmine salicylate solution (Eserine; Sigma Millipore) was given to mice under brief isoflurane anesthesia. The concentration of eserine solution was increased from 0.1%, 0.15%, and 0.2% based on the pupil size. Basal oculomotor performances were described in gains. For the actual learning protocol, we adopted OKR and it consisted of five training sessions (10 min per session), six checkup points, and 24 hr of consolidation period. During the training sessions, the mice were visually stimulated with a sinusoidally rotating drum (±5 degrees). Completely trained mice were placed back into their home cage, which was stored in the dark condition until the last gain checkup points.

Gain Analysis

Gains acquired from basal oculomotor performances and OKR learning were calculated as the ratio of evoked eye movements to the movement of the screen or turn table as visual or vestibular stimuli, respectively. A custom-built LabView (National Instrument) analysis tool was used for all the calculations

Statistical analysis

Graph plotting and statistical analysis were performed using GraphPad Prism (GraphPad Software Inc, CA, USA). The hypothesis was tested by using paired and unpaired t-tests between sample pairs using GraphPad Prism. Statistical significance was set at p < 0.05. Asterisks denoted in the graph indicate the statistical significance. * indicates p < 0.05, ** < 0.01, *** < 0.001, and **** < 0.0001. The test name and statistical values are presented in each figure legend.